Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions for the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
%matplotlib inline
from pyquil.parameters import Parameter, quil_sin, quil_cos
from pyquil.quilbase import DefGate
```

Output:
```
Available frameworks:
Forest SDK
Qiskit
D-Wave Ocean
```

# State preparation



As explained in Maria Schuld's guest lecture, to prepare a state in a particular encoding is done by rotating vectors, in other words to map vectors into others, which is the basic idea of kernel learning. This procedure gives rise to interesting kernels, without the need of specifying the mapping function.


Let us work with an easy training data set of two vectors, $S = \{(\begin{bmatrix}0 \\ 1\end{bmatrix}, 0), (\begin{bmatrix}\frac{1}{2}\sqrt{2} \\ \frac{1}{2}\sqrt{2}\end{bmatrix}, 1)\}$. Let's have a test instance $\begin{bmatrix}1 \\ 0\end{bmatrix}$. We will build the interference circuit for this.

**Exercise 1** (1 point).

Create a circuit in your preferred framework that works on four qubits: ancilla, index, data, and class. Put the ancilla and index qubits into a uniform superposition. Verify for yourself that this also encodes the test data in the coefficients of the $|0000\rangle$, $|0100\rangle$ states as well as in the $|0010\rangle$, $| 0110\rangle$ states.

Insert identity operators for the data_qubit and the class_qubit, so these are initialized for amplitude measurements. 

Place your solution in an object called `circuit`. 

```python
ancilla_qubit = 0
index_qubit = 1
data_qubit = 2
class_qubit = 3


### BEGIN SOLUTION

# Forest version

circuit = Program()
# Put the ancilla and the index qubits into uniform superposition
circuit += H(ancilla_qubit)
circuit += H(index_qubit)
# Nothing to be done to prepare the test state, but the I operations
# are needed to ensure amplitude measurements run over all 4-qubit states
# (we could put these in assignment_helper.py)
circuit += I(data_qubit)
circuit += I(class_qubit)

# QISKIT VERSION
q = QuantumRegister(4)
c = ClassicalRegister(4)
circuit = QuantumCircuit(q, c)
# Put the ancilla and the index qubits into uniform superposition
# qiskit does not need the identity operators

circuit.h(q[ancilla_qubit])
circuit.h(q[index_qubit])

# Nothing to be done to prepare the test state.

### END SOLUTION


# Verify for yourself that this preparation is also sufficient to represent the test vector
```

Output:
```
[0.5+0.j 0.5+0.j 0.5+0.j 0.5+0.j 0. +0.j 0. +0.j 0. +0.j 0. +0.j 0. +0.j
 0. +0.j 0. +0.j 0. +0.j 0. +0.j 0. +0.j 0. +0.j 0. +0.j]
```

<!-- grade_id: exercise1, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0.5 +0.j, 0.5+0.j, 0.5 +0.j, 0.5+0.j, 0.+0.j, 0. +0.j, 0.+0.j,
       0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j]))
```

**Exercise 2** (1 point).

Extend the circuit to prepare the first training instance, by pushing the feature values into the coefficients of the $|0001\rangle$ and $|0101\rangle$, without affecting any of the other coefficients.

```python
### BEGIN SOLUTION

# Forest version
# Prepare the first training vector
# We can prepare this with a Toffoli
circuit += CCNOT(ancilla_qubit, index_qubit, data_qubit)
# Flip the index qubit > moves the first training vector to the 
# |0> state of the index qubit
circuit += X(index_qubit)


# QISKIT VERSION
# Prepare the first training vector
# We can prepare this with a Toffoli
circuit.ccx(ancilla_qubit, index_qubit, data_qubit)
# Flip the index qubit > moves the first training vector to the 
# |0> state of the index qubit
circuit.x(index_qubit) 

###END SOLUTION
```

<!-- grade_id: exercise2, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0.5 +0.j, 0. +0.j, 0.5 +0.j, 0.5+0.j, 0.+0.j, 0.5+0.j, 0.+0.j,
       0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j]))
```

**Exercise 3** (2 points).

As a final step in the preparation, extend the circuit to push the feature values of the second training instance into the coefficients of the $|0011\rangle$ and $|0111\rangle$ states.

Remember that, as explained in the lecture notebook, loading vectors with amplitude encoding requires to apply a double controlled rotation by an angle determined by the coordinates of the vector.


```python
### BEGIN SOLUTION

# Forest version
# This can be done in qiskit style (see below) or with a constructed double controlled RY
# We choose the latter and define the rotation

alpha = Parameter('alpha')
ccry = np.array([
    [1, 0, 0, 0, 0, 0, 0, 0],
    [0, 1, 0, 0, 0, 0, 0, 0],
    [0, 0, 1, 0, 0, 0, 0, 0],
    [0, 0, 0, 1, 0, 0, 0, 0],
    [0, 0, 0, 0, 1, 0, 0, 0],
    [0, 0, 0, 0, 0, 1, 0, 0],
    [0, 0, 0, 0, 0, 0, quil_cos(alpha / 2), -quil_sin(alpha / 2)],
    [0, 0, 0, 0, 0, 0, quil_sin(alpha / 2), quil_cos(alpha / 2)]
])
ccry_definition = DefGate('CCRY', ccry,[alpha])
CCRY = ccry_definition.get_constructor()
# Now apply the CCRY
circuit += ccry_definition
circuit += CCRY(np.pi/2)(ancilla_qubit, index_qubit, data_qubit)


# QISKIT VERSION
# see Bloch diagram in lecture on movement of relevant states
circuit.ccx(ancilla_qubit, index_qubit, data_qubit)
circuit.ry(-np.pi/4, data_qubit)
circuit.ccx(ancilla_qubit, index_qubit, data_qubit)
circuit.ry(np.pi/4, data_qubit)

###END SOLUTION
```

<!-- grade_id: exercise3, points: 2, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0.5 +0.j,  0. +0.j,  0.5 +0.j,  np.sqrt(2)/4 + 0.j,  0.+0.j,  0.5+0.j,
       0.+0.j,  np.sqrt(2)/4 +0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j, 0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j]))
```

**Exercise 4** (1 point). 

Finish the state preparation circuit by a conditional flip of the class qubit, such that the coefficients of the second training vector and those of the copy of the test vector have the class bit equal to 1, without affecting any of the other coefficients.


It is very useful at this point, for testing purposes, to add code which shows the state vectors, as done in the lecture notebook


```python
### BEGIN SOLUTION

# Forest version
circuit += CNOT(index_qubit, class_qubit)

# QISKIT VERSION
circuit.cx(q[index_qubit], q[class_qubit])

### END SOLUTION
```

<!-- grade_id: exercise4, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([ 0.5 +0.j,  0.  +0.j,  0. +0.j,  0.+0.j, 0.  +0.j,  0.5 +0.j,  0. +0.j,  0.+0., 0.  +0.j,  0.  +0.j,  0.5+0.j,  np.sqrt(2)/4 +0j, 0.  +0.j,  0.  +0.j,  0. +0.j,  np.sqrt(2)/4 +0.j]))
```

# Interference as a kernel

**Exercise 5** (1 point). 

At this point, our state is:

$|\psi\rangle = \frac{1}{\sqrt{2}}\sum_{m=0}^{1}|y_m\rangle|m\rangle|\psi_{x^m}\rangle|0\rangle + |y_m\rangle|m\rangle|\psi_{\tilde{x}}\rangle|1\rangle$

Apply the Hadamard gate on the ancilla to produce the interference and create: 

$ |\psi\rangle = \frac{1}{2\sqrt{2}}\sum_{m=0}^{1}|y_m\rangle|m\rangle(|\psi_{x^m}\rangle+|\psi_{\tilde{x}}\rangle)|0\rangle+|y_m\rangle|m\rangle(|\psi_{x^m}\rangle-|\psi_{\tilde{x}}\rangle)|1\rangle$


```python
 ### BEGIN SOLUTION
    
# Forest version
circuit += H(0)

#QISKIT VERSION
circuit.h(q[ancilla_qubit])

### END SOLUTION
```

<!-- grade_id: exercise5, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([  np.sqrt(2)/4+0.j , np.sqrt(2)/4+0.j , 0.        +0.j , 0.        +0.j,
  np.sqrt(2)/4+0.j ,-np.sqrt(2)/4+0.j , 0.        +0.j , 0.        +0.j,
  0.        +0.j,  0.        +0.j , (np.sqrt(2)+1)/4+0.j , (np.sqrt(2)-1)/4+0.j,
  0.        +0.j , 0.        +0.j , 0.25      +0.j ,-0.25      +0.j]))
```

If we measure the ancilla, the outcome probability of observing 0 will be $\frac{1}{4N}\sum_{i=1}^N |x_t + x_i|^2$. Performing post-selection on the 0 outcome, we can calculate the kernel and the probability of the test instance belonging to either class.