Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions for the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
%matplotlib inline
```

# Quantum approximate optimization algorithm

QAOA is a shallow-circuit variational algorithm that is easy to understand if you already grasped quantum annealing. It is, in fact, just a particular type of a larger family of algorithms called variational quantum eigensolvers. Approximating the adiabatic pathway as QAOA does is just one option of how to find the eigenvalues of a system.

Even then, QAOA has many moving elements. Let us import some handy packages and define some functions that we are going to use:

```python
import itertools
from functools import partial, reduce
from scipy.optimize import minimize
np.set_printoptions(precision=3, suppress=True)

# Functions useful if you're using Qiskit
def pauli_x(qubit, coeff):
    eye = np.eye((n_qubits))
    return PauliOp(Pauli(np.zeros(n_qubits), eye[qubit]), coeff)

def pauli_z(qubit, coeff):
    eye = np.eye((n_qubits))
    return PauliOp(Pauli(eye[qubit], np.zeros(n_qubits)), coeff)

def product_pauli_z(q1, q2, coeff):
    eye = np.eye((n_qubits))
    return PauliOp(Pauli(eye[q1], np.zeros(n_qubits)) * Pauli(eye[q2], np.zeros(n_qubits)), coeff)
```

**Exercise 1** (2 points). Define a mixing Hamiltonian on two qubits. Store it in an object called `Hm`. If you're doing the assignement in PyQuil, remember that the Hamiltonian should be a list of PauliTerms (and not a PauliSum) in order to be exponentiated easily.

```python
n_qubits = 2
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise1, points: 2, grade: true, solution: false -->
```python
if isinstance(Hm, SummedOp):
    assert np.alltrue(Hm.to_matrix() == np.array([[0., 1., 1., 0.],
                                                  [1., 0., 0., 1.],
                                                  [1., 0., 0., 1.],
                                                  [0., 1., 1., 0.]]))
elif isinstance(Hm, list):
    assert len(Hm) == n_qubits
    assert all([isinstance(Hm[i], PauliTerm) for i in range(n_qubits)])
    assert all([Hm[i].compact_str() == '(1+0j)*X{}'.format(i) for i in range(n_qubits)])
else:
    raise ValueError("Unknown type for Hamiltonian!")
```

**Exercise 2** (2 points). Define the cost Hamiltonian $H_c = -\sigma^Z_1\sigma^Z_2-0.5\sigma^Z_1$.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise2, points: 2, grade: true, solution: false -->
```python
if isinstance(Hc, SummedOp):
    assert np.alltrue(Hc.to_matrix() == np.diag([-1.5, 1.5, 0.5, -0.5]))
    
elif isinstance(Hc, list):
    assert len(Hc) == 2
    assert all([isinstance(Hc[i], PauliTerm) for i in range(n_qubits)])
    assert Hc[0].compact_str() == '(-1+0j)*Z0Z1'
    assert Hc[1].compact_str() == '(-0.5+0j)*Z0'
else:
    raise ValueError("Unknown type for Hamiltonian!")    
```

**Exercise 3** (2 points). We know that the ground state of the mixer Hamiltonian is the uniform superposition. Create a circuit `circuit_init` that will contain this initial state.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise3, points: 2, grade: true, solution: false -->
```python
amplitudes = init_state.to_matrix()
assert np.allclose(amplitudes, np.array([0.5, 0.5, 0.5, 0.5]))
```

We set $p=2$ and initialize the $\beta_i$ and $\gamma_i$ parameters. 

```python
p = 2
beta = np.random.uniform(0, np.pi*2, p)
gamma = np.random.uniform(0, np.pi*2, p)
```

The next step is to create the complete variational circuit, made of $e^{-\beta H}$ and $e^{-\gamma H}$. We will use a function `create_circuit` that takes `gamma` and `beta` as argument, and the state preparation circuit.

```python
def evolve(hamiltonian, angle):
    return (angle * hamiltonian).exp_i()

def create_circuit(init_state, beta, gamma):
    if isinstance(init_state, qiskit.aqua.operators.StateFn):
        circuit = ComposedOp([evolve(Hm, beta[i]) @ evolve(Hc, gamma[i]) for i in range(p)] + [init_state])
    elif isinstance(circuit_init, pyquil.quil.Program):
        exp_Hm = []
        exp_Hc = []
        for term in Hm:
            exp_Hm.append(exponential_map(term))
        for term in Hc:
            exp_Hc.append(exponential_map(term))
        circuit = Program()
        circuit += circuit_init
        for i in range(p):
            for term_exp_Hm in exp_Hm:
                circuit += term_exp_Hm(-beta[i])
            for term_exp_Hc in exp_Hc:
                circuit += term_exp_Hc(-gamma[i])
    return circuit
```

Finally, we need a function `evaluate_circuit` to compute the average energy of the circuit, i.e. compute $\langle\psi(\beta, \gamma)|H_c|\psi(\beta, \gamma)\rangle$ where $|\psi(\beta, \gamma)\rangle$ is the circuit built above. This function should take a unique argument `beta_gamma` (concatenation of the lists `beta` and `gamma`) in order to be used directly by optimizers, and return a real value corresponding to the expectancy of $H_c$.

```python
def evaluate_circuit(beta_gamma):
    n = len(beta_gamma)//2
    circuit = create_circuit(init_state, beta_gamma[:n], beta_gamma[n:])
    if isinstance(circuit, qiskit.aqua.operators.OperatorBase):
        meas = OperatorStateFn(Hc, is_measurement=True)
        circuit_meas = PauliTrotterEvolution(trotter_mode=Suzuki(order=3, reps=1)).convert(meas @ circuit)
        circuit_expectation = PauliExpectation().convert(circuit_meas)
        
        sampler = CircuitSampler(backend=Aer.get_backend('statevector_simulator'))
        sampled_trotter_exp_op = sampler.convert(circuit_expectation)
        return sampled_trotter_exp_op.eval().real
    elif isinstance(circuit, pyquil.quil.Program):
        qvm = pyquil.api.QVMConnection(endpoint=fc.sync_endpoint, compiler_endpoint=fc.compiler_endpoint)
        return np.real(qvm.pauli_expectation(circuit, sum(Hc)))
```

**Exercise 4** (2 points). The $p$ parameter defines the number of steps in the Trotterization. The real question from here is how we optimize the $\beta_i$ and $\gamma_i$ parameters. If we can find a method that makes fewer evaluations to arrive at the same result, that is a win, since we have to execute fewer loops on the quantum computer. Try various methods for minimizing the evaluate function. We used L-BFGS-B before. Try another one and write the outcome in an object called `result`. You will see that the number of function evaluation (`nfev`) differs and so does the function value.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise4, points: 2, grade: true, solution: false -->
```python
result2 = minimize(evaluate_circuit, np.concatenate([beta, gamma]), method='L-BFGS-B')
import scipy
assert isinstance(result, scipy.optimize.optimize.OptimizeResult)
assert result2.nfev != result.nfev
print("Function evaluations: %d versus %d" % (result2.nfev, result.nfev))
print("Function values: %f versus %f" % (result2.fun, result.fun))
```

If the circuit had an infinite capacity, you should obtain a minimum of $-1.5$ (minimum value of the Hamiltonian we defined above). However, with $p=2$, you might have a bigger value. Try increasing $p$ to see the effect on the minimum reached by the circuit.

Variational circuits are actually very similar to neural networks: we do some kind of gradient descent over a parameter space. Automatic differentiation of neural networks has been a major step in scaling them up and we can expect that similar techniques in variational quantum circuit can be useful. This is exactly what projects like [QuantumFlow](https://github.com/rigetti/quantumflow) and [PennyLane](https://pennylane.ai/) try.