Before you begin, execute this cell to import numpy and packages from the D-Wave Ocean suite, and all necessary functions for the gate-model framework you are going to use, whether that is the Forest SDK or Qiskit. In the case of Forest SDK, it also starts the qvm and quilc servers.

```python
%run -i "assignment_helper.py"
%matplotlib inline
```

# State preparation

Preparing a state in a particular encoding can, in itself, give rise to interesting kernels. This is what talked about in the lecture on the interference kernel, and Maria Schuld's guest lecture expands on the idea. Let us work with an easy data set of two vectors, $S = \{(\begin{bmatrix}0 \\ 1\end{bmatrix}, 0), (\begin{bmatrix}\sqrt{2}/2 \\ \sqrt{2}/2\end{bmatrix}, 1)\}$. Let's have a test instance $\begin{bmatrix}1 \\ 0\end{bmatrix}$. We will build the interference circuit for this.

**Exercise 1** (1 point). Create a circuit in your preferred framework that works on four qubits: ancilla, index, data, and class. Put the ancilla and index qubits into a uniform superposition. Entangle the test instance with the ground state of the ancilla. Put an identity gate on the class qubit. Place your solution in an object called `circuit`. 

```python
ancilla_qubit = 0
index_qubit = 1
data_qubit = 2
class_qubit = 3
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise1, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0. +0.j, 0.5+0.j, 0. +0.j, 0.5+0.j, 0.5+0.j, 0. +0.j, 0.5+0.j,
       0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j]))
```

**Exercise 2** (1 point). Extend the circuit to prepare the first training instance and entangle it with the excited state of the ancilla and ground state of the index qubit.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise2, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([0. +0.j, 0. +0.j, 0. +0.j, 0.5+0.j, 0.5+0.j, 0.5+0.j, 0.5+0.j,
       0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j, 0. +0.j]))
```

**Exercise 3** (2 points). Extend the circuit to prepare the second training instance and entangle it with the excited state of the ancilla and the excited of the index qubit.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise3, points: 2, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([ 0. +0.j,  0. +0.j,  0. +0.j,  0.5+0.j,  0.5+0.j,  0.5+0.j,
       -0.5+0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j, 0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j]))
```

**Exercise 4** (1 point). Finish the state preparation circuit by flipping the class qubit conditioned on the index qubit.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise4, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([ 0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0.5+0.j,  0.5+0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0. +0.j,  0.5+0.j, 0. +0.j,  0. +0.j, -0.5+0.j,  0. +0.j]))
```

# Interference as a kernel

**Exercise 5** (1 point). At this point, our state is $\frac{1}{2\sqrt{2}}\sum_{i=0}^1|0\rangle|i\rangle(|x_t\rangle+|x_i\rangle)|y_i\rangle+|1\rangle|i\rangle(|x_t\rangle-|x_i\rangle)|y_i\rangle$, where $|x_t\rangle$ is the encoded test instance and $|x_i\rangle$ is a training instance. Apply the Hadamard gate on the ancilla to apply the interference.

```python
###
### YOUR CODE HERE
###

```

<!-- grade_id: exercise5, points: 1, grade: true, solution: false -->
```python
amplitudes = get_amplitudes(circuit)
assert np.allclose(amplitudes, np.array([ 0.        +0.j,  0.        +0.j,  0.        +0.j,  0.        +0.j,
        0.70710678+0.j,  0.        +0.j,  0.        +0.j,  0.        +0.j,  0.        +0.j,  0.        +0.j,  0.35355339+0.j, -0.35355339+0.j,
        0.        +0.j,  0.        +0.j, -0.35355339+0.j, -0.35355339+0.j]))
```

If we measure the ancilla, the outcome probability of observing 0 will be $\frac{1}{4N}\sum_{i=1}^N |x_t + x_i|^2$. Performing post-selection on the 0 outcome, we can calculate the kernel and the probability of the test instance belonging to either class.